#  Metrologie 2
De programmeer opdrachten voor Metrologie 2 bestaan uit Jupyter werkbladen. Deze werkbladen moeten voor de deadline (zie studiewijzer) ingeleverd zijn op Blackboard. 

De werkbladen worden door de computer nagekeken. Daarom is het belangrijk om je aan bepaalde afspraken te houden:
* de functie templates moeten gebruikt worden
* waar staat 
``` # YOUR CODE HERE ``` tot en met ```raise NotImplementedError()``` moet vervangen worden door je eigen code.
* let op eventuele return regels, daar moet er 1 van per functie overblijven.
* test je eigen code voordat je het inlevert!

## Jupyter
Jupyter notebooks zijn bestanden waarbij code en tekst door elkaar lopen. Door een 'cell' (blokje tekst of code) als code aan te merken wordt de cell uitvoerbaar. Eventuele output van de code in de cell komt in de ruimte onder de cell te staan. Als voorbeeld onderstaande cell:

```python
# cell om modules te importeren. Nodig voor interne werking document
import numpy as np
import numpy.testing as t
import matplotlib.pyplot as plt

```

```python
som = 0
N = 10
for i in range(N):
    som += i
print('De som van de getallen van 0 tot {} is {}.'.format(N, som))
```

In dit werkblad gaan jullie werken aan integratie algorithmes. De functies die jullie schrijven kan je buiten de jupyter omgeving in bijvoorbeeld IDLE of Spyder direct gebruiken. Met CTRL-a kan de inhoud van een cell geselecteerd worden, om daarna met CTRL-c en CTRL-v de data te kopieren en te plakken.

### Oefenopdracht 1
Schrijf een functie met de naam `oppervlak` die voor een rechthoek met lengte $l$ en breedte $b$ de oppervlakte uitrekent. 

<!-- grade_id: oefenopdracht, points: 0, grade: false, solution: true -->
```python
def oppervlak(l,b):
    ''' geeft het oppervlak voor een rechthoek met lengte l en breedte b.'''
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: oefenopdracht-tests, points: 1, grade: true, solution: false -->
```python
# TEST FUNCTIES
t.assert_allclose(oppervlak(1,2),2)
t.assert_allclose(oppervlak(3,3),9)
t.assert_allclose(oppervlak(3,0),0)
```

### Opdracht 1
Op bladzijde 123 en 124 van het boek Computational Physics van M. Newman staat een implementatie van een lineaire fit aan willekeurige data. In het boek Research Methods for Science van M.P. Marder staat ook een uitwerking van een lineaire fit, met een mogelijkheid de onzekerheid in de parameters van de fit te bepalen.

Door een grotere dataset op te delen in kleine stukken en van deze (lineaire) delen de vergelijking op te stellen is, afhankelijk van de variatie in de dataset een willekeurige functie als meerdere rechte stukken te integreren. Deze rechte lijnen zijn zelfs met de hand uit te rekenen, daarna is het een simpele optelsom van de deel-oppervlakken.

Maak een functie (`linfit`) die als input een lijst met gemeten x-waarden y-waarden krijgt:
$$ y = a\cdot x + b. $$
Daarnaast geeft de functie de onzekerheid in de parameters $a$ en $b$ terug. Daarnaast wordt ook de 't' waarde teruggegeven. Het functie template is in de cel hieronder weergegeven.

<!-- grade_id: linfit, points: 0, grade: false, solution: true -->
```python
def linfit(x,y):
    ''' functie die voor een dataset (x,y) waarden de helling (a), offset (b), onzekerheid in a en b
    (delta_a, delta_b) en t-test resultaat (t) teruggeeft. '''
    
    if len(x) != len(y): # x,y data moeten even lang zijn!
        raise ValueError('''Lengte van de input data is niet gelijk!
        aantal elementen in x: {}
        aantal elementen in y: {}'''.format(len(x), len(y)))
    
    N = len(x) # aantal datapunten in x,y
    
    # YOUR CODE HERE
    raise NotImplementedError()
    return (a,b,delta_a,delta_b,t)
```

Om bijvoorbeeld 100 gemeten punten (x,y) weer te geven met de lineaire fit kan onderstaande code gebruikt worden.

De y-data wordt gegenereerd door de formule
$$y = (3 \pm 0.3) + (4 \pm 0.4)x,$$
met een meetonnauwkeurigheid 
de lineaire fit zal dus niet exact overeen komen met de 'gemeten' data.

```python
# voorbeeld gebruik
def randomdata(a,b):
    return a + (b-a)*np.random.random()

err_a = 0.3
err_b = 0.4
x = range(10)
y = [(3+randomdata(-err_a,err_a)) + (4+randomdata(-err_b,err_b))*i for i in x]
yerr = [randomdata(-err_a, err_a) + randomdata(-err_b,err_b)*i for i in x]
try:
        a,b,delta_a, delta_b,tt = linfit(x,y)
except ValueError as e:
        print(e)

# zie ook https://hhs-tn.github.io        
def func(x,a,b):
    return a*x + b

# vul ymax/ymin in met behulp van sigma maal onzekerheid in fit parameters
p_opt=np.array([a,b])
fit_error=np.array([delta_a, delta_b])
sigma_plot =3
ymax = func(x, *(p_opt+fit_error*sigma_plot))
ymin = func(x, *(p_opt-fit_error*sigma_plot))
ydata = func(x, *p_opt)

for ii in range(len(fit_error)):
    rc = p_opt[0] + sigma_plot*fit_error[0]
    fit_error[0] = -fit_error[0]
    for jj in range(len(fit_error)):
        offset = p_opt[1] + sigma_plot*fit_error[1]
        fit_error[1] = -fit_error[1]
        yerr = func(x, rc, offset)
        maxTrue = yerr > ymax
        minTrue = yerr < ymin
        ymax[yerr > ymax] = yerr[yerr > ymax]
        ymin[yerr < ymin] = yerr[yerr < ymin]
        
        
# grafische weergave        
plt.errorbar(x,y,yerr=yerr,fmt='k*',label='data')
plt.plot(x,ydata, label='fit')
plt.plot(x,ymin, label='$3\sigma')
plt.plot(x,ymax, label='$3\sigma')
plt.show()

print('''De gevonden vergelijking is:
 y = {b:4.5f} + {a:4.5f}x
met als onnauwkeurigheid in de parameters
  a = {a:4.5f} +/- {da:4.5f}
  b = {b:4.5f} +/- {db:4.5f}'''.format(b=b,a=a,db=3*delta_b,da=3*delta_a))
```

<!-- grade_id: linfit-basis, points: 1, grade: true, solution: false -->
```python
# TEST FUNCTIES 
x = range(-10, 10)
y = [3+4*i+0.1*i**2 for i in x]
a,b,delta_a, delta_b, tt = linfit(x,y)
t.assert_allclose(a,3.9)
t.assert_allclose(b,6.3)
t.assert_allclose(a,3.9)
```

<!-- grade_id: linfit-basis2, points: 1, grade: true, solution: false -->
```python
# TEST FUNCTIES 
## geen offset
x = range(-10,10)
y = [3+4*i for i in x]
a,b,delta_a, delta_b, tt = linfit(x,y)
t.assert_allclose(a,4)
t.assert_allclose(b,3)
t.assert_allclose(delta_a, 0)
t.assert_allclose(delta_b, 0)
```

<!-- grade_id: linfit-basis3, points: 1, grade: true, solution: false -->
```python
# HIDDEN TEST FUNCTIES
```

```python

```